# Programming Assignment 1 - IBAN

The assignments are released on Mondays. You have **5 days** to solve and upload them (=> **upload by Friday** at the latest).
**At least 80%** of all the reachable assignment points have to be achieved by the end of the semester.

Please use the variable names as given below.

### Topic
The International Bank Account Number (IBAN) is an internationally agreed system of identifying bank accounts. An IBAN number starts with a <font color="blue">country code</font> of two letters for the country, followed by two <font color="magenta">check digits</font>, the <font color="green">bank code</font> and finally the <font color="red">account number</font>.<br><br>
<font color="blue">DE</font><font color="magenta">29</font><font color="green">10010010</font><font color="red">0987654321</font><br>
<font color="blue">country code </font>|<font color="magenta"> check digits </font>|<font color="green"> bank code </font>|<font color="red"> account number</font><br><br>
To calculate the check digits (<font color="magenta">??</font>) of <font color="blue">DE</font><font color="magenta">??</font><font color="green">10010010</font><font color="red">0987654321</font> proceed as follows (see also https://www.ibantest.com/en/how-is-the-iban-check-digit-calculated).

### Exercise 1 (2 Points)

#### Task:
The country code has to be converted into numbers.  
For A = 10, B = 11 etc. up to Z = 35. For example "<font color="blue">DE</font>" becomes the number "<font color="blue">13</font>" and "<font color="blue">14</font>". Find the values of the first and second letter.  
Cou can use the function `ord()`. For example, `ord(“A”)` returns 65, `ord("B")` returns 66, etc.  
The value of the variable `first_letter` should have value "13" <b>(as a string!)</b>. Also set the variable `second_letter` accordingly.

<!-- grade_id: cell-265de131ca34cd06, points: 0, grade: false, solution: true -->
```python
iban = "DE29100100100987654321"

first_letter = iban[0]
second_letter = iban[1]
first_letter = ord(first_letter)
second_letter = ord(second_letter)

first_letter = str(first_letter - 55)
second_letter = str(second_letter - 55)
print(first_letter)
print(second_letter)



```

Output:
```
13
14
```

<!-- grade_id: cell-54261f08e1c526f0, points: 2, grade: true, solution: false -->
```python
# Test section

assert str(type(first_letter)) == "<class 'str'>"
assert first_letter == "13"

assert str(type(second_letter)) == "<class 'str'>"
assert second_letter == "14"
```

### Exercise 2 (1 Point)

#### Task:
Use <font color="green">bank code</font> and <font color="red">account number</font> and append the converted <font color="blue">country code</font> along with two zeros (00).  
iban_part = <font color="green">10010010</font><font color="red">0987654321</font><font color="blue">1314</font>00  
<b>Don't just copy paste the number</b>, rather use `first_letter` and `second_letter` you've defined in <b>Exercise 1</b>

<!-- grade_id: cell-5da2a3886685db6a, points: 0, grade: false, solution: true -->
```python
iban_part = iban[4:]
country_code = first_letter + second_letter
iban_part = iban_part + first_letter + second_letter + "00"
print(iban_part)


```

Output:
```
100100100987654321131400
```

<!-- grade_id: cell-26805dec7bf70d1d, points: 1, grade: true, solution: false -->
```python
# Test section

assert str(type(iban_part)) == "<class 'str'>"
assert iban_part == "100100100987654321131400"
```

### Exercise 3 (1 Point)

#### Task:
Calculate the resulting number modulo 97, i.e. <font color="green">10010010</font><font color="red">0987654321</font><font color="blue">1314</font>00 mod 97 = 69  
The check digit is '98 minus this result', i.e.  
98 - 69 = <font color="magenta">29</font>  
Please write the result into the variable `iban_mod`.

<!-- grade_id: cell-1610cffaebb8dcf2, points: 0, grade: false, solution: true -->
```python
iban_mod = 98 - (int(iban_part)%97) 
```

<!-- grade_id: cell-4058a6fc3b2d8f77, points: 1, grade: true, solution: false -->
```python
# Test section

assert str(type(iban_mod)) == "<class 'int'>"
assert iban_mod == 29
```

### Exercise 4 (1 Point)

#### Task:
Write a comparison with an if-statement with the comparison operator <i>==</i> if the <font color="magenta">check digits</font> are correct (i.e. if they are the same as the iban_mod).  
`is_check_true` shall be `True` if they are the same or `False` in case they are not.

<!-- grade_id: cell-2916ec167f407b6a, points: 0, grade: false, solution: true -->
```python
is_check_true = True 
check_digit = int(iban[2:4])
if check_digit == iban_mod:
 is_check_true = True
else:
 is_check_digit = False

```

<!-- grade_id: cell-be109e654dbbf0fd, points: 1, grade: true, solution: false -->
```python
# Test section

assert str(type(is_check_true)) == "<class 'bool'>"
assert is_check_true == True
```

### Exercise 5 (5 Points)

#### Task:
Put everything together in a function named `iban_verification`.  
If the IBAN is correct it should return `True`; if the IBAN is incorrect it should return `False`.  

#### Example:
Here, we check the validity of the IBAN from "Welthungerhilfe":
Calling `iban_verification("DE15370501980000001115")` returns `True`, because `"DE15370501980000001115"` is a valid IBAN.

<!-- grade_id: cell-94969a54c4cc2d9e, points: 0, grade: false, solution: true -->
```python

def iban_verification(iban):
 first_letter = iban[0]
 second_letter = iban[1]
 first_letter = ord(first_letter)
 second_letter = ord(second_letter)

 first_letter = str(first_letter - 55)
 second_letter = str(second_letter - 55)
 iban_part = iban[4:]
 iban_part = iban_part + first_letter + second_letter + "00"


 iban_mod = 98 - (int(iban_part)%97) 
 is_check_true = False 
 check_digit = int(iban[2:4])
 if check_digit == iban_mod:
  is_check_true = True
 else:
  is_check_digit = False
 return is_check_true
```

```python
# For testing purposes

iban_verification("DE15370501980000001115")
# Shall return
# True
```

Output:
```
True
```

<!-- grade_id: cell-53e4c4d28542a4f5, points: 5, grade: true, solution: false -->
```python
# Test section

# Testing for True
assert iban_verification("DE15370501980000001115") == True
assert iban_verification("DE31370501980000033225") == True
assert iban_verification("ES9121000418450200051332") == True
assert iban_verification("FI2112345600000785") == True
assert iban_verification("AT611904300234573201") == True
# Testing for False
assert iban_verification("AT15370501980000001115") == False
assert iban_verification("DE37370501980000033225") == False
assert iban_verification("ES9121000418450200051662") == False
assert iban_verification("FI2189347600000785") == False
assert iban_verification("DE611904300234573201") == False

```

```python

```